<center>
<img src="logo.png" height="900"> 
</center>


#  Анализируем чеки

В этом задании мы будем работать с покупками и чеками. Смотреть за корреляциями в покупках довольно полезно.

> В 1992 году группа по консалтингу в области ритейла компании Teradata под руководством Томаса Блишока провела исследование 1.2 миллиона транзакций в 25 магазинах для ритейлера Osco Drug (Drug Store — формат разнокалиберных магазинов у дома). После анализа всех этих транзакций самым сильным правилом получилось «Между 17:00 и 19:00 чаще всего пиво и подгузники покупают вместе». К сожалению, такое правило показалось руководству Osco Drug настолько контринтуитивным, что ставить подгузники на полках рядом с пивом они не стали. Хотя объяснение паре пиво-подгузники вполне себе нашлось: когда оба члена молодой семьи возвращались с работы домой (как раз часам к 5 вечера), жены обычно отправляли мужей за подгузниками в ближайший магазин. И мужья, не долго думая, совмещали приятное с полезным — покупали подгузники по заданию жены и пиво для собственного вечернего времяпрепровождения.

Для работы будем использовать датасет о продуктовых корзинах: https://www.kaggle.com/heeraldedhia/groceries-dataset

```python
import numpy as np
import pandas as pd

import scipy.stats as sts
import matplotlib.pyplot as plt
import seaborn as sns

plt.style.use('ggplot')  # стиль для графиков
%matplotlib inline
```

Подружаем данные и смотрим как они выглядят.

```python
df = pd.read_csv('groceries.csv', sep=',')
df.columns = ['id', 'fielddate', 'product']
print(df.shape)
df.head(20)
```

Output:
```
(38765, 3)
      id   fielddate                    product
0   1808  21-07-2015             tropical fruit
1   2552  05-01-2015                 whole milk
2   2300  19-09-2015                  pip fruit
3   1187  12-12-2015           other vegetables
4   3037  01-02-2015                 whole milk
5   4941  14-02-2015                 rolls/buns
6   4501  08-05-2015           other vegetables
7   3803  23-12-2015                 pot plants
8   2762  20-03-2015                 whole milk
9   4119  12-02-2015             tropical fruit
10  1340  24-02-2015               citrus fruit
11  2193  14-04-2015                       beef
12  1997  21-07-2015                frankfurter
13  4546  03-09-2015                    chicken
14  4736  21-07-2015                     butter
15  1959  30-03-2015      fruit/vegetable juice
16  1974  03-05-2015  packaged fruit/vegetables
17  2421  02-09-2015                  chocolate
18  1513  03-08-2015              specialty bar
19  1905  07-07-2015         
```

## 1. Корреляции

Для начала поработаем с корреляциями в данных. 

__а)__ Какой товар покупался чаще всего? Сохраните название этого товара в переменную `product_name`.

<!-- grade_id: cell-9d268bdaaa451d51, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz

# will the code be with you

# Формат ответа: строка, пример: 'pip fruit'
# Вокруг ответа не должно быть никаких надписей array(['pip fruit']) или Index('pip fruit')
# Это должна быть именно строка, а не строка в массиве
product_name = str(df['product'].value_counts().index[0])


# your code here

```

<!-- grade_id: cell-2fd1cfdcf5400ee3, points: 1, grade: true, solution: false -->
```python
# проверка, что задание решено корректно
assert len(product_name) == 10

# Аналогичные тесты скрыты от вас
```

__б)__ Сколько всего уникальных заказов было сделано? Сохраните число заказов в переменную `n_cnt`.

<!-- grade_id: cell-4c3ce54e9c52ae14, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

# Формат ответа: целое число, пример: 5555
n_cnt = df['id'].dropna().unique().size
n_cnt


# your code here

```

Output:
```
3898
```

<!-- grade_id: cell-608432c41bdcd71d, points: 1, grade: true, solution: false -->
```python
# проверка, что задание решено корректно
assert n_cnt > 3800
assert n_cnt < 4000

# Аналогичные тесты скрыты от вас
```

В таблице выше в каждой строке записана информация о покупке конкретного товара. Давайте сделаем табличку размера "число товаров" на "число покупок", чтобы понимать какие товары покупались вместе, а какие нет. 

> Обратите внимание, то здесь задание немного упрощено. Вообще говоря, нам нужно делать агрегацию по паре `fielddate, id`, если мы хотим изучать чеки по-честному. Но мы делаем её только по `id` для того, чтобы не усложнять задание. В качестве необязательного дополнения вы можете после сдачи задания переделать код так, чтобы дата тоже учитывалась при расчётах. 

```python
sparse_sales = pd.pivot_table(df, 
               values='fielddate', 
               index='id', 
               columns='product', 
               fill_value=0, aggfunc='count')

sparse_sales.head()
```

Output:
```
product  Instant food products  UHT-milk  abrasive cleaner  artif. sweetener  \
id                                                                             
1000                         0         0                 0                 0   
1001                         0         0                 0                 0   
1002                         0         0                 0                 0   
1003                         0         0                 0                 0   
1004                         0         0                 0                 0   

product  baby cosmetics  bags  baking powder  bathroom cleaner  beef  berries  \
id                                                                              
1000                  0     0              0                 0     0        0   
1001                  0     0              0                 0     1        0   
1002                  0     0              0                 0     0        0   
1003                  0     0     
```

В нашей матрице огромное число нулей. Обычно такие матрицы называют разряжеными. Мы занимаем нулями кучу свободной памяти, которую мы могли бы не занимать, если бы хранили данные [в ином виде.](https://cmdlinetips.com/2018/03/sparse-matrices-in-python-with-scipy/)

__в)__ Постройте матрицу корреляций Пирсона. Для этого используйте метод таблицы `.corr`.

<!-- grade_id: cell-a8f15b8ba3a946c8, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

# Формат ответа: таблица pd.DataFrame размера (167, 167)
sales_correlation = sparse_sales.corr()
sales_correlation

# your code here

```

Output:
```
product                Instant food products  UHT-milk  abrasive cleaner  \
product                                                                    
Instant food products               1.000000 -0.006936         -0.009420   
UHT-milk                           -0.006936  1.000000          0.013806   
abrasive cleaner                   -0.009420  0.013806          1.000000   
artif. sweetener                   -0.010825  0.006105         -0.006523   
baby cosmetics                     -0.003470 -0.007877         -0.002091   
...                                      ...       ...               ...   
white bread                         0.016640 -0.011550          0.010777   
white wine                          0.002853  0.009990          0.000107   
whole milk                          0.015981  0.028747          0.003558   
yogurt                              0.005071  0.006505          0.008360   
zwieback                            0.018221  0.028753         -0.009420   

product    
```

Какие продукты сильнее всего коррелируют с яйцами, `domestic eggs` (их чаще всего покупают вместе)?  Сохраните название самого скоррелированного продукта в переменную `top_1`.

<!-- grade_id: cell-cf20442762e0f104, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

# Формат ответа: строка, пример: 'pip fruit'
# Вокруг ответа не должно быть никаких надписей array(['pip fruit']) или Index('pip fruit')
# Это должна быть именно строка, а не строка в массиве
top_1 = str(sales_correlation['domestic eggs'].sort_values().index[-2])
top_1


# your code here

```

Output:
```
'meat spreads'
```

Какие продукты "мешают" купить яйца, то есть отрицательно коррелируют с их покупкой? Сохраните название продукта с самой большой отрицательной корреляцией в переменную `bottom_1`.

<!-- grade_id: cell-fbd202e1529e31c7, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

# Формат ответа: строка, пример: 'pip fruit'
# Вокруг ответа не должно быть никаких надписей array(['pip fruit']) или Index('pip fruit')
# Это должна быть именно строка, а не строка в массиве
bottom_1 = str(sales_correlation['domestic eggs'].sort_values().index[0])
bottom_1

# your code here

```

Output:
```
'pet care'
```

<!-- grade_id: cell-d070d4a1a7bde720, points: 1, grade: true, solution: false -->
```python
# проверка, что задание решено корректно
assert len(bottom_1) == 8
assert len(top_1) == 12

# Аналогичные тесты скрыты от вас
```

Напишите код, который выводит самые коррелируемые товары для случайного продукта из списка `unique_products`.

<!-- grade_id: cell-50a7be49564df467, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

unique_products = df['product'].unique()
t=np.random.choice(unique_products)
str(sales_correlation[t].sort_values().index[-2])

# your code here

```

Output:
```
'nut snack'
```

__г)__ Какие два продукта коррелируют сильнее всего? Положите их название в лист `answer`

<!-- grade_id: cell-dd462f4db3d9a2f7, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

# Формат ответа: массив из строк, пример: ['soups', 'apple', 'iphone 18']
unique_products = df['product'].unique()
list1=[]
list2=[]
answer=[]
d=np.array([])
for prod in unique_products:
    list1.append(sales_correlation[prod].sort_values().index[-2])
    list2.append(sales_correlation[prod].sort_values()[-2])
correl=0
for i in range(167):
    if list2[i]>=correl:
        correl=list2[i]
        k=i
prod1=list1[k]
prod2=sales_correlation[prod1].sort_values().index[-2]
answer.append(prod1)
answer.append(prod2)
sparse_sales
# your code here

```

Output:
```
product  Instant food products  UHT-milk  abrasive cleaner  artif. sweetener  \
id                                                                             
1000                         0         0                 0                 0   
1001                         0         0                 0                 0   
1002                         0         0                 0                 0   
1003                         0         0                 0                 0   
1004                         0         0                 0                 0   
...                        ...       ...               ...               ...   
4996                         0         0                 0                 0   
4997                         0         0                 0                 0   
4998                         0         0                 0                 0   
4999                         0         0                 0                 0   
5000                         0         0
```

<!-- grade_id: cell-894ff9bec07f24e0, points: 1, grade: true, solution: false -->
```python
# проверка, что задание решено корректно
assert 'soups' in answer

# Аналогичные тесты скрыты от вас
```

Конечно же, корреляция — это [не единственный способ искать](https://habr.com/ru/company/ods/blog/353502/) между покупками ассоциативные правила.

## 2. Зависимость. 

В лекции мы с вами сказали, что события $A$ и $B$ называются независимыми, если $P(AB) = P(A)\cdot P(B)$. Отталкиваясь от этого определения, можно ввести другую характеристику, которая показывает, насколько продукты зависят друг от друга, а именно __поддержку (lift).__ 

$$
lift = \frac{P(AB)}{P(A)\cdot P(B)}
$$

Эта метрика описывает отношение зависимости товаров к их независимости. Если оказалось, что `lift = 1`, это означает, что покупка товара $A$ не зависит от покупки товара $B$. Если `lift > 1`, то это означает, что вероятность встретить оба товара в чеке, $P(AB)$ высокая, то есть товары покупают вместе. Если `lift < 1`, это означает, что товары, наоборот, очень часто покупают по-отдельности. 

__д)__ Посчитайте значение нашей метрики для яиц и молока (`'whole milk', 'domestic eggs'`). Запишите получившиеся значение метрики в переменную `answer`.

> Вам аккуратно нужно сделать три среза по условию  `>= 1`. Там, где пара надо делать срез так, чтобы оба товара дали `True`. Сделать это в одну строку вам поможет метод `.all(axis=1)`. Частоты можно получить методом `.mean()`, так как python думает, что `False` - это ноль, а `True` - это единица.

<!-- grade_id: cell-db191a336be19f97, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

# Формат ответа: действительное чиссло, пример: 3.1418281828
a='whole milk'
b='domestic eggs'
a_and_b=sparse_sales[(sparse_sales[a]>=1) & (sparse_sales[b]>=1)]
Pa=sparse_sales[(sparse_sales[a]>=1)]
cPa=float(Pa[a].count())
p=float(sparse_sales.shape[0])
Pb=sparse_sales[(sparse_sales[b]>=1)]
cPb=float(Pb[b].count())
supp=float(a_and_b[a].count())
lift=(supp/p)/((cPa/p)*(cPb/p))
answer = float(lift)
answer

# your code here

```

Output:
```
1.152241691425711
```

<!-- grade_id: cell-c2f789696293a0b3, points: 1, grade: true, solution: false -->
```python
# проверка, что задание решено корректно
assert answer < 3
assert answer > 1

# Аналогичные тесты скрыты от вас
```

__е)__ Посчитайте значение метрики для всех пар продуктов из датасета. Сохраните значения в словарик `dict`. В качестве ключа используете кортеж из пары продуктов. Чтобы удобнее было перебрать все сочетания, используйте `combinations` из модуля `itertools`.

Чтобы при подсчётах не возникало деления на ноль, добавьте к знаменателю маленькое число, например `1e-10`.

<!-- grade_id: cell-e512719bd3dbbd34, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you
from itertools import combinations
dict={}
s=tuple(combinations(sparse_sales.columns, 2))
def val(a:str,b:str):
    a_and_b=sparse_sales[(sparse_sales[a]>=1) & (sparse_sales[b]>=1)]
    Pa=sparse_sales[(sparse_sales[a]>=1)]
    cPa=float(Pa[a].count())
    p=float(sparse_sales.shape[0])
    Pb=sparse_sales[(sparse_sales[b]>=1)]
    cPb=float(Pb[b].count())
    supp=float(a_and_b[a].count())
    lift=(supp/p)/((cPa/p)*(cPb/p))
    answer1 = float(lift)
    return answer1

for i in range(len(s)):
    a=s[i][0]
    b=s[i][1]
    dict[s[i]]=val(a,b)
print(dict)
#dict[0]
#for val in s:
    
# your code here

```

Output:
```
{('Instant food products', 'UHT-milk'): 0.8492374727668844, ('Instant food products', 'abrasive cleaner'): 0.0, ('Instant food products', 'artif. sweetener'): 0.0, ('Instant food products', 'baby cosmetics'): 0.0, ('Instant food products', 'bags'): 0.0, ('Instant food products', 'baking powder'): 0.5369146005509642, ('Instant food products', 'bathroom cleaner'): 0.0, ('Instant food products', 'beef'): 1.2547210300429186, ('Instant food products', 'berries'): 0.6266881028938908, ('Instant food products', 'beverages'): 1.0738292011019284, ('Instant food products', 'bottled beer'): 1.1544964997307485, ('Instant food products', 'bottled water'): 0.8579031612645057, ('Instant food products', 'brandy'): 0.0, ('Instant food products', 'brown bread'): 0.98062893081761, ('Instant food products', 'butter'): 1.0542258282623393, ('Instant food products', 'butter milk'): 0.7703557312252965, ('Instant food products', 'cake bar'): 2.9198501872659173, ('Instant food products', 'candles'): 0.0, ('Insta
```

Сколько пар продуктов покупали вместе хотя бы раз? Запишите ответ в переменную `answer`.

<!-- grade_id: cell-097e70120aa5a596, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

# Формат ответа: целое число, пример: 5555
answer=0
for key,val1 in dict.items():
    if val1!=0.0:
        answer+=1      
        

# your code here

```

Output:
```
9824
```

Для какой пары продуктов метрика $lift$ оказалась самой большой? 

<!-- grade_id: cell-e0e772420c99a1f4, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you
max_lift=0.0
for key,val1 in dict.items():
    if val1>max_lift:
        max_lift=val1
        key_max_lift=key
print(key_max_lift)

# your code here

```

Output:
```
('preservation products', 'soups')
```

Сколько раз эти продукты встретились в выборке? Как думаете адеватно ли делать выводы по такому объёму данных? 

<!-- grade_id: cell-cfbf5b5deb321f36, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you
a1='preservation products'
b1='soups'
a1andb1=sparse_sales[(sparse_sales[a1]>=1) &( sparse_sales[b1]>=1)]
meetings=a1andb1.shape[0]
meetings # нет, неадекватно.
# your code here

```

Output:
```
1
```

Для какой пары продуктов метрика оказывается самой маленькой? 

<!-- grade_id: cell-0c88c82cc7bdef09, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you
min_lift=2000
for key,val1 in dict.items():
    if val1<min_lift:
        min_lift=val1
        key_min_lift=key
print(key_min_lift,min_lift)#для всех продуктов , где метрика 0 , то есть которые никогда не были куплены вместе, код можно написать по -другому, но и так видно)
# your code here

```

Output:
```
('Instant food products', 'abrasive cleaner') 0.0
```

<!-- grade_id: cell-b7f41317d840457e, points: 1, grade: true, solution: false -->
```python
# проверка, что задание решено корректно

assert answer < 10000
assert answer > 9000

# Аналогичные тесты скрыты от вас
```

## 3. Неоцениваемые задания

Выше мы увидели, что некоторые продукты встречаются в выборке очень редко. Понятное дело, что по ним у нас не получится построить хорошее ассоциативное правило. Попробуйте повторить расчёт той же метрики, но с условием что продукт покупали больше 10 раз. Изучите самые покупаемые вместе продукты и самые непокупаемые вместе продукты. Насколько сильно список отличается от полученного в предыдущем задании? 

<!-- grade_id: cell-aa30f3933e22e20d, points: 0, grade: false, solution: true -->
```python
### ╰( ͡° ͜ʖ ͡° )つ▬▬ι═══════  bzzzzzzzzzz
# will the code be with you

# your code here

```

Иногда в чеках пытаются искать __продукты-якоря.__ То есть продукты, которые являются основными. Например: айфон - основной продукт, наушники и чехол - дополнения к нему. Подумайте как можно попытаться найти такие продукты на основе простых метрик, основанных на подсчёте условных вероятностей.

<center>
<img src="https://pp.userapi.com/c638028/v638028181/52e5e/1X-dkzNN1hk.jpg" width="400"> 
</center>